```python
# Initialize Otter
import otter
grader = otter.Notebook("Assignment03_COMM187Fall2024.ipynb")
```

# Coding Assignment #3
## COMM187: Data Science in Communication Research, Fall 2024

**!!! Please make sure to run the first cell before running auto-grader !!!**

### Question 1

**1 POINT**

Make a list of 550 random integers ranging from 0 to 1,000 and save it in a new variable `newlist`.

***Tip:*** Use `random` library as shown in Coding Lab #2.

```python
### Write your code below (in place of ...)
...
```

```python
grader.check("q1")
```

### Question 2: 

**1 POINT**

Fetch the value at the first position in `newlist`, and store it in a new variable `firstval`.

***Tip:*** Refer to the section on "Indexing Lists" in Coding Lab #2. Remember that indexing in Python starts at 0!

```python
### Write your code below (in place of ...)
firstval = ...
```

```python
grader.check("q2")
```

### Question 3

**1.5 POINT**

Fetch the value at the 49th position in `newlist`, and store it in a new variable `val49`.

***Tip:*** Refer to the section on "Indexing Lists" in Coding Lab #2. Remember that indexing in Python starts at 0!

```python
### Write your code below (in place of ...)
val49 = ...
```

```python
grader.check("q3")
```

### Question 4

**1.5 POINT**

Calculate the average of the values at the 90th, 167th, 210th, and 435th positions in `newlist`, and store the output in a new variable `valmean`.

```python
### Write your code below (in place of ...)
valmean = ...
```

```python
grader.check("q4")
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Assignment #3 for COMM187: Data Science in Communication Research. Once finished answering questions, first **SAVE** then download this .ipynb file. Submit the file as instructed on Canvas and Gradescope. **ONLY** submit the .ipynb file, not the pdf.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True)
```

Tests for q1 (points: 1)
```
>>> max(newlist) <= 1000
True
```
```
>>> min(newlist) >= 0
True
```
```
>>> len(newlist) == 550
True
```


Tests for q2 (points: 1)
```
>>> firstval == newlist[0]
True
```
```
>>> firstval != newlist[1]
True
```


Tests for q3 (points: 1.5)
```
>>> val49 == newlist[48]
True
```
```
>>> val49 != newlist[49]
True
```
```
>>> val49 != newlist[50]
True
```


Tests for q4 (points: 1.5)
```
>>> valmean == (newlist[89] + newlist[166] + newlist[209] + newlist[434]) / 4
True
```
