```python
# Initialize Otter
import otter
grader = otter.Notebook("Assignment02_COMM187Fall2024.ipynb")
```

# Coding Assignment #2
## COMM187: Data Science in Communication Research, Spring 2024

**!!! Please make sure to run the first cell before running auto-grader !!!**

### Question 1

**0.5 POINT**

Let us practice some boolean operators. 

We have two boolean variables below, `bool_01` and `bool_02`. 

Perform the following operations and store the output in the new variables as instructed. Refer to this week's lab for help.

| Variable Name | Operation |
| :------------ | :-------- |
| `bool_and`    | "and" operation |
| `bool_or`    | "or" operation |
| `bool_xor`    | "exclusive or" operation |

```python
### Write your code below (in place of ...)
bool_01 = False
bool_02 = True

bool_and = ...
bool_or = ...
bool_xor = ...
```

```python
grader.check("q1")
```

### Question 2: 

**1 POINT**

You have just downloaded data from Reddit for four individuals: Ace, Buck, and Cliff, and Dime. 

For each of them, you have the number of reddit posts they have made since Jan 1, 2024. These are stored in the variables `posts_Ace` for Ace, `posts_Buck` for Buck, `posts_Cliff` for Cliff, and `posts_Dime` for Dime. 

Your task: test whether each of them made more than 1,000 posts. 

Make four new boolean variables, `ismorethan1000_Ace`  for Ace, `ismorethan1000_Buck` for Buck, `ismorethan1000_Cliff` for Cliff, and `ismorethan1000_Dime` for Dime. For each of these variables, their value should be `True` if the number of posts made by that individual is greater than 1,000. If not, their value should be `False`.

***Tip***: Pay attention to the data type of `posts_Ace`, `posts_Buck`, `posts_Cliff`, and `posts_Dime`. Do they need to be typecast into another data type first?

***Tip***: Make sure you are writing `1000` when you mean to say "one thousand"! If you write `1,000`, Python may think you are saying `1` and `000`, not one thousand. 

```python
### Write your code below (in place of ...)
posts_Ace = "904"
posts_Buck = "1010"
posts_Cliff = "1000"
posts_Dime = "870"

ismorethan1000_Ace = ...
ismorethan1000_Buck = ...
ismorethan1000_Cliff = ...
ismorethan1000_Dime = ...
```

```python
grader.check("q2")
```

### Question 3

#### A.

**1 POINT**

A communication researcher is analyzing the time spent by one individual on *Instagram*. \
They collected data on the time spent (in minutes) on Instagram for 7 consecutive days. \
The times spent on each day platforms by this individual are stored in variables `ig_day_01`, `ig_day_02`, `ig_day_03`, `ig_day_04`, `ig_day_05`, `ig_day_06`, and `ig_day_07`, as shown below. 

Write a Python code that makes a **list** with the times spent on Instagram in sequence from day 1 to day 7. The list should be stored in a new variable `ig_list`.

```python
### Write your code below (in place of ...)
ig_day_01 = 100
ig_day_02 = 120
ig_day_03 = 25
ig_day_04 = 105
ig_day_05 = 165
ig_day_06 = 110
ig_day_07 = 115
...
```

```python
grader.check("q3a")
```

#### B.

**0.5 POINT**

Immediately following the 7 consecutive days, the communication researcher continues to observe the time spent by the same individual on Instagram for three more consecutive days.

On day 8, this individual spends 80 minutes on Instagram. \
On day 9, they spend 190 minutes on Instagram. \
On day 10, they spend 140 minutes on Instagram.

Make a new list which just has these three data points in sequence, i.e., time spent on day 8 comes first, then day 9, then day 10. 

Store your output in a new list called `ig_list_later`.

```python
### Write your code below (in place of ...)
ig_list_later = ...
```

```python
grader.check("q3b")
```

#### C.

**0.5 POINT**

Make an updated list which combines the two lists you have just made, `ig_list` and `ig_list_later`, such that it contains the time spent by this individual on Instagram in sequence from day 1 to day 10.

Store your output in a new list called `ig_list_updated`.

**Tip:** You can add two lists to join them using the `+` operator!

```python
### Write your code below (in place of ...)
ig_list_updated = ...
```

```python
grader.check("q3c")
```

#### D.

**0.5 POINT**

Now, the communication researcher wants to calculate the minimum and the maximum daily time spent by this individual over 10 days.

Store the minimum daily time in `min_ig` and maximum daily time in `max_ig`.

**Tip:** Use the `min()` and `max()` functions from Coding Lab #1!

```python
### Write your code below (in place of ...)
...
```

```python
grader.check("q3d")
```

#### E

**1 POINT**

Finally, the communication researcher wants to make a guess about what this individual's daily activity might look like for the next 100 days. For that, the researcher creates a list of random integers with a minimum of `min_ig` and a maximum of `max_ig`.

Make a new variable `ig_random_100` and store this new random list in this variable.

**Tip:** Use the `random.sample()` function from Coding Lab #2! Think carefully about the different inputs to this function!

```python
### Write your code below (in place of ...)
...
```

```python
grader.check("q3e")
```

#### BONUS

**0.5 POINT**

As you may have noticed, in `ig_list_updated`, all values are divisible by 5. However, this might not be the case in `ig_random_100`.

For a bonus 0.5 point, create a new list `ig_random_20_by5` with 20 random integers ranging from `min_ig` to `max_ig` which are divisible by 5.

**Tip:** Think about the "step" value in the range function!

```python
### Write your code below (in place of ...)
...
```

```python
grader.check("bonus")
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Assignment #2 for COMM187: Data Science in Communication Research. Once finished answering questions, first **SAVE** then download this .ipynb file. Submit the file as instructed on Canvas and Gradescope. **ONLY** submit the .ipynb file, not the pdf.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True)
```

Tests for bonus (points: 0.5)
```
>>> max(ig_random_20_by5) <= 190
True
```
```
>>> min(ig_random_20_by5) >= 25
True
```
```
>>> len(ig_random_20_by5) == 20
True
```
```
>>> all((i % 5 == 0 for i in ig_random_20_by5))
True
```


Tests for q1 (points: 0.5)
```
>>> bool_and == False
True
```
```
>>> bool_or == True
True
```
```
>>> bool_xor == True
True
```


Tests for q2 (points: 1)
```
>>> ismorethan1000_Ace == False
True
```
```
>>> ismorethan1000_Buck == True
True
```
```
>>> ismorethan1000_Cliff == False
True
```
```
>>> ismorethan1000_Dime == False
True
```


Tests for q3a (points: 1)
```
>>> len(ig_list) == 7
True
```
```
>>> ig_list == [100, 120, 25, 105, 165, 110, 115]
True
```


Tests for q3b (points: 0.5)
```
>>> len(ig_list_later) == 3
True
```
```
>>> ig_list_later == [80, 190, 140]
True
```


Tests for q3c (points: 0.5)
```
>>> len(ig_list_updated) == 10
True
```
```
>>> ig_list_updated == [100, 120, 25, 105, 165, 110, 115, 80, 190, 140]
True
```


Tests for q3d (points: 0.5)
```
>>> min_ig == 25
True
```
```
>>> max_ig == 190
True
```


Tests for q3e (points: 1)
```
>>> max(ig_random_100) <= 190
True
```
```
>>> min(ig_random_100) >= 25
True
```
```
>>> len(ig_random_100) == 100
True
```
