```python
# Initialize Otter
import otter
grader = otter.Notebook("Assignment04_COMM187Fall2024.ipynb")
```

# Coding Assignment #4
## COMM187: Data Science in Communication Research, Fall 2024

**!!! Please make sure to run the first cell before running auto-grader !!!**

### Question 1

**1 POINT**

Imagine you are conducting a research survey and you have a variable, `response_time`, that records the *number of seconds* it took each respondent to answer a question. Below, you are presented with the `response_time` for one particular respondent.

Write a Python conditional statement (`if`, `else`, and `elif`) to categorize responses as "Quick", "Average", or "Slow" based on the following conditions:

 - **Quick** if the response time is less than 5 seconds.
 - **Average** if the response time is between 5 and 15 seconds (inclusive).
 - **Slow** if the response time is greater than 15 seconds.

**Tip:** Use `if` followed by `elif` followed by `else` statements!

```python
### Write your code below (in place of ...)

response_time = 12

...
```

```python
grader.check("q1")
```

### Question 2: 

**2 POINTS**

In a research project, you might need to create a dataset representing different groups of survey respondents. 

Let's create an array to represent the participation status of 200 respondents in a survey. \
A value of **1** indicates **participation**, while **0** indicates **non-participation**.

Use `numpy` to create:
 - An numpy array of size 200 where all values are 1 (indicating all participants took part in the survey) and store it in variable `participants`.
 - An numpy array of size 200 where all values are 0 (indicating none of the participants took part in the survey) and store it in variable `non_participants`.


```python
### Write your code below (in place of ...)
import numpy
...
```

```python
grader.check("q2")
```

### Question 3

**2 POINTS**

Suppose you are analyzing the number of *likes* and *shares* on the social media posts made by an individual, collected at different times throughout the day. The data for like are stored in the numpy array `likes` and the shares are stored in numpy array `shares`.

As the data scientist for this project, you have deviced a new measure called "total engagement" which is the sum of likes and shares for a given social media post. 

For this individual, write the Python code which calculates the "total engagement", which is the sum of likes and shares, and store it in a new variable called `total_engagement`.

**Tip:** To add two numpy arrays item by item, you can simply use the `+` operator. Refer to Coding Lab #4 for more help!

```python
### Write your code below (in place of ...)
post_likes = numpy.array([449, 471, 779, 626, 370, 254, 627, 298, 711, 315, 478, 565])
post_shares = numpy.array([75, 26, 54, 36,  2, 45, 31,  1, 27, 71, 67, 32])
...
```

```python
grader.check("q3")
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Assignment #4 for COMM187: Data Science in Communication Research. Once finished answering questions, first **SAVE** then download this .ipynb file. Submit the file as instructed on Canvas and Gradescope. **ONLY** submit the .ipynb file, not the pdf or zip file.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True)
```

Tests for q1 (points: 1)
```
>>> response_category == 'Average'
True
```


Tests for q2 (points: 2)
```
>>> participants.size == 200
True
```
```
>>> non_participants.size == 200
True
```
```
>>> all(participants == 1)
True
```
```
>>> all(non_participants == 0)
True
```


Tests for q3 (points: 2)
```
>>> numpy.array_equal(total_engagement, post_likes + post_shares)
True
```
