```python
# Initialize Otter
import otter
grader = otter.Notebook("Assignment5_COMM187Spring2024.ipynb")
```

# Coding Assignment #5
## COMM187: Data Science in Communication Research, Spring 2024

**!!! Please make sure to run the first cell before running auto-grader !!!**

<h3><span style="color:green"> Finish assignment on DataCamp </span> </h3>

This week's questions are based on the assigned chapter on DataCamp. If you have not yet signed into this class's DataCamp group, please use this invite link and make an account using your UCSB email ID: https://www.datacamp.com/groups/shared_links/0d032623fe95677c03dd5d41331db87feeb0738725bb7ae390f6d9ee17f2bed8 

This week, you have been assigned the chapter **Intermediate Python: Dictionaries and Pandas** on DataCamp. \
Finish this assigned chapter before proceeding with this assignment. You will need the skills taught in that chapter to solve this week's coding assignment.

**Question 1: Creating Dictionary**

In a study about the messaging strategy of an election campaign, the number of messages sent out before the election date were analyzed. 

Create a dictionary named `messages_dict` that contains the following key:value pairs, representing the number of messages delivered through different media:

 - "Email: 120
 - "SMS": 340
 - "Call": 50
 - "Instagram": 35
 - "Twitter": 20

```python
import numpy as np
import pandas as pd
```

```python
### Write your code below (in place of ...)
messages_dict = {"Email": 120, "SMS": 340, "Call": 50, "Instagram": 35, "Twitter": 20} # SOLUTION
```

```python
grader.check("q1")
```

**Question 2. Indexing and Filtering pandas DataFrames** 

In a study about effects of sharing posts on social media on mental health, you have the data for participants and their corresponding communication media usage in a pandas DataFrame named `socialmedia_df` with the following keys and values:

 - "Participant": a list of participant IDs
 - "SharedPost": a list of number of posts shared by each participant
 - "SocialMedia": a list of media used by each participant

Run the following cell to load the data in to a DataFramed named `socialmedia_df` and to display it.

```python
socialmedia_df = pd.read_csv('./data/assign05.csv')
display(socialmedia_df)
```

Output:
```
Participant  SharedPosts SocialMedia
0       ppt_000           65    Facebook
1       ppt_001           55    Facebook
2       ppt_002           28      TikTok
3       ppt_003           51   Instagram
4       ppt_004           18      TikTok
..          ...          ...         ...
116     ppt_116           16     Twitter
117     ppt_117           61   Instagram
118     ppt_118           98     Twitter
119     ppt_119           72   Instagram
120     ppt_120            5     Twitter

[121 rows x 3 columns]
```

Now, answer the following questions:

**a.** From `socialmedia_df`, select the first 20 rows and store the resulting DataFrame in a variable named `df_filtered_iloc`.

Hint: Use `.iloc[]`!

```python
### Write your code below (in place of ...)
# BEGIN SOLUTION
df_filtered_iloc = socialmedia_df.iloc[0:20]
# END SOLUTION
```

```python
grader.check("q2a")
```

**b.** Set the row index of `socialmedia_df` to the column "Participant"

Use the function `.set_index()` and take help from [this tutorial](https://www.geeksforgeeks.org/python-pandas-dataframe-set_index/).

```python
### Write your code below (in place of ...)
# BEGIN SOLUTION
socialmedia_df.set_index("Participant", inplace = True)
# END SOLUTION
```

```python
grader.check("q2b")
```

**c.** Select the value at row index "ppt_081" and column "SocialMedia" from `socialmedia_df` and store it in a new variable `cell_value`.

Hint: Use `.loc[]`

The autograder for this question is hidden. It will directly be graded on Gradescope.

```python
### Write your code below (in place of ...)
# BEGIN SOLUTION
cell_value = socialmedia_df.loc["ppt_081", "SocialMedia"]
# END SOLUTION
```

**d.** From `socialmedia_df`, select the subset of rows where "SharedPost" value is greater than or equal to 50.

Store the resulting DataFrame in a variable named `df_50plus`.

```python
### Write your code below (in place of ...)
# BEGIN SOLUTION
df_50plus = socialmedia_df[socialmedia_df["SharedPosts"] >= 50]
# END SOLUTION
```

```python
grader.check("q2d")
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Assignment #4 for COMM187: Data Science in Communication Research. Once finished answering questions, first **SAVE** then download this .ipynb file. Submit the file as instructed on Canvas and Gradescope. **ONLY** submit the .ipynb file, not the zip file.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True)
```

Tests for q1 (points: 2)
```
>>> all([k in ['Email', 'SMS', 'Call', 'Instagram', 'Twitter'] for k in messages_dict.keys()])
True
```
```
>>> messages_dict['Email'] == 120
True
```
```
>>> messages_dict['SMS'] == 340
True
```
```
>>> messages_dict['Call'] == 50
True
```
```
>>> messages_dict['Instagram'] == 35
True
```
```
>>> messages_dict['Twitter'] == 20
True
```


Tests for q2a (points: 3)
```
>>> df_filtered_iloc.shape == (20, 3)
True
```
```
>>> [df_filtered_iloc.SharedPosts.min(), df_filtered_iloc.SharedPosts.max()] == [1, 90]
True
```
```
>>> all(df_filtered_iloc.SocialMedia.value_counts() == [7, 6, 5, 2])
True
```
```
>>> df_filtered_iloc.iloc[-1, 0] == 'ppt_019'
True
```


Tests for q2b (points: 3)
```
>>> all([p in socialmedia_df.index for p in ['ppt_000', 'ppt_050', 'ppt_120']])
True
```


Tests for q2c (points: 3)
```
>>> cell_value == 'TikTok'
True
```


Tests for q2d (points: 5)
```
>>> df_50plus.shape[0] == 68
True
```
```
>>> df_50plus[df_50plus.SharedPosts < 50].shape[0] == 0
True
```
```
>>> all(df_50plus.SocialMedia.value_counts() == [21, 18, 15, 14])
True
```
