```python
# Initialize Otter
import otter
grader = otter.Notebook("Assignment1_COMM187Spring2024.ipynb")
```

# Assignment #1
## COMM187: Data Science in Communication Research, Spring 2024

**Question 1:** A communication researcher is analyzing the average time spent by individuals on different social media platforms. They collected data on the time spent (in minutes) on six platforms: *Facebook*, *Instagram*, *Twitter/X*, *Tiktok*, *Snapchat*, and *YouTube*. The times spent on these platforms by a particular individual are stored in variables `facebook_time`, `instagram_time`, `twitterx_time`, `tiktok_time`, `snapchat_time`, and `youtube_time`, as shown below. 

Write a Python code that calculates the *average time spent on these platforms* by this individual, and assign it to a new variable `average_time`. Use the variables provided and make sure to include the correct number of platforms in your calculation.

```python
### Write your code below (in place of ...)
facebook_time = 40
instagram_time = 115
twitterx_time = 15
tiktok_time = 100
snapchat_time = 20
youtube_time = 35
# BEGIN SOLUTION
average_time = (facebook_time + instagram_time + twitterx_time + tiktok_time + snapchat_time + youtube_time)/6
# END SOLUTION
```

```python
grader.check("q1")
```

**Question 2:** A social media influencer wants to calculate the percentage increase in their followers on Instagram over the last month. They had `initial_followers` at the beginning of the month and `final_followers` at the end of the month. Write a Python code that calculates the percentage increase in their followers using these variables, and assigns this value to a new variable `percentage_increase`. 

***Tip:*** Subtract the initial number of followers from the final number before dividing by the initial number of followers, and then multiply by 100 to get the percentage.

```python
### Write your code below (in place of ...)
initial_followers = 39024 
final_followers = 46784
# BEGIN SOLUTION
percentage_increase = ((final_followers - initial_followers)/initial_followers)*100
# END SOLUTION
```

```python
grader.check("q2")
```

**Question 3:** 

**a)** A digital marketing researcher is studying the effectiveness of a social media ad campaign for a product. They want to calculate the "purchase rate", which is the percentage of people who made an online purchase after clicking on an ad. If `total_clicks` people clicked on the ad and `total_purchases` people made an online purchase, write a Python expression that calculates the purchase rate using these variables, and assigns this value to a new variable `purchase_rate`. 

***Tip:*** Divide the number of purchases by the number of clicks and then multiply by 100 to get the percentage.

```python
### Write your code below (in place of ...)
total_clicks = 1239099
total_purchases = 5912
# BEGIN SOLUTION
purchase_rate = (total_purchases/total_clicks)*100
# END SOLUTION
```

```python
grader.check("q3_1")
```

**b)** The digital marketing researcher now also wants to predict future purchases through this digital ad campaign. They first decide to predict the number of clicks in the future. According to their analysis, the number of clicks grow by 1.2 times every week.

If this pace of growth keeps up, what will be the total number of clicks by the end of the next 4 weeks? Calculate the number of clicks in the next 4 weeks, and assign it to a new variable `total_clicks_future`.

***Tip:*** In one week, the number of clicks will go from `total_clicks` to `total_clicks * 1.2`. Similarly, in two week, the number of clicks will go up to `total_clicks * 1.2 * 1.2`. Think about which operation you can use to calculate the total number of clicks in 4 weeks. \
***Another tip:*** Number of clicks can be counted as positive integers. They cannot be negative. They cannot have decimal points. If you get a number with decimal points, round them to the nearest integer.

```python
### Write your code below (in place of ...)
# BEGIN SOLUTION
total_clicks_future = round(total_clicks*(1.2**4))
# END SOLUTION
```

```python
grader.check("q3_2")
```

**c)** Now that the digital marketing researcher has calculated the predicted number of clicks in the next 4 weeks in `total_clicks_future`, they would like to predict the total number of purchases by the end of the next 4 weeks. 

Calculate the total number of predicted purchases by the end of 4 weeks, which is `total_clicks_future` times the `purchase_rate` calculate above, divided by 100, and assign this value to a new variable `predicted_purchases`.

***Tip:*** Just like number of clicks, number of purchases are positive integers too. They cannot be negative. They cannot have decimal points. If you get a number with decimal points, round them to the nearest integer.

```python
### Write your code below (in place of ...)
# BEGIN SOLUTION
predicted_purchases = round((total_clicks_future*purchase_rate)/100)
# END SOLUTION
```

```python
grader.check("q3_3")
```

**d)** Finally, the digital marketing researcher would like to calculate the predicted increase in the number of purchases in 4 weeks.

Calculate the predicted increase in purchases in 4 weeks by calculating the difference between `predicted_purchases` and `total_purchases`, and assigning it to a new variable `predicted_purchase_increase`.

```python
### Write your code below (in place of ...)
# BEGIN SOLUTION
predicted_purchase_increase = predicted_purchases - total_purchases
# END SOLUTION
```

```python
grader.check("q3_4")
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

This is a demo assignment for COMM187: Data Science in Communication Research. Once finished answering questions, first **SAVE** then download this .ipynb file. Submit the file as instructed on Canvas and Gradescope.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True)
```

Tests for q1 (points: 1)
```
>>> average_time == (40 + 115 + 15 + 100 + 20 + 35) / 6
True
```


Tests for q2 (points: 2)
```
>>> percentage_increase == (46784 - 39024) / 39024 * 100
True
```


Tests for q3_1 (points: 2)
```
>>> purchase_rate == 5912 / 1239099 * 100
True
```


Tests for q3_2 (points: 2)
```
>>> total_clicks_future == 2569396
True
```
```
>>> total_clicks_future == 1239099 * 1.2 ** 4 if total_clicks_future != 2569396 else True
True
```


Tests for q3_3 (points: 2)
```
>>> predicted_purchases == 12259
True
```
```
>>> predicted_purchases == 2569396 * (5912 / 1239099 * 100) / 100 if predicted_purchases != 12259 else True
True
```


Tests for q3_4 (points: 1)
```
>>> predicted_purchase_increase == 6347
True
```
```
>>> predicted_purchases == 2569396 * (5912 / 1239099 * 100) / 100 if predicted_purchases != 12259 else True
True
```
