```python
# Initialize Otter
import otter
grader = otter.Notebook("Assignment2Solution_COMM187Spring2024.ipynb")
```

# Coding Assignment #2
## COMM187: Data Science in Communication Research, Spring 2024

**!!! Please make sure to run the first cell before running auto-grader !!!**

**Question 1:** Let us practice some boolean operators. 

We have two boolean variables below, `bool_01` and `bool_02`. 

Perform the following operations and store the output in the new variables as instructed. Refer to this week's lab for help.

| Variable Name | Operation |
| :------------ | :-------- |
| `bool_and`    | "and" operation |
| `bool_or`    | "or" operation |
| `bool_xor`    | "exclusive or" operation |

```python
### Write your code below (in place of ...)
bool_01 = True
bool_02 = False

bool_and = bool_01 & bool_02 # SOLUTION
bool_or = bool_01 | bool_02 # SOLUTION
bool_xor = bool_01 ^ bool_02 # SOLUTION
```

```python
grader.check("q1")
```

**Question 2:** You have just downloaded data from Twitter/X on three individuals, let's call them individual 1 , individual 2, and individual 3. For each of them, you have the number of tweets/posts they have made since Jan 1, 2024. These are stored in the variables `tweets_indv_01` for individual 1, `tweets_indv_02` for individual 2, and `tweets_indv_03` for individual 3. 

Your task: test whether each of them made more than 500 tweets. 

Make three new boolean variables, `ismorethan500_indv_01`  for individual 1, `ismorethan500_indv_02` for individual 2, and `ismorethan500_indv_03` for individual 3. For each of these variables, their value should be True if the number of tweet for that individual is greater than 500. If not, their value should be False.

Hint: Pay attention to the data type of `tweets_indv_01`, `tweets_indv_02`, and `tweets_indv_03`. Do they need to be typecast into another data type first?

```python
### Write your code below (in place of ...)
tweets_indv_01 = "234"
tweets_indv_02 = "500"
tweets_indv_03 = "502"

ismorethan500_indv_01 = int(tweets_indv_01) > 500 # SOLUTION
ismorethan500_indv_02 = int(tweets_indv_02) > 500 # SOLUTION
ismorethan500_indv_03 = int(tweets_indv_03) > 500 # SOLUTION
```

```python
grader.check("q2")
```

**Question 3:** 

**a.** A communication researcher is analyzing the time spent by one individual on *Instagram*. They collected data on the time spent (in minutes) on Instagram for 7 consecutive days. The times spent on each day platforms by this individual are stored in variables `ig_day_01`, `ig_day_02`, `ig_day_03`, `ig_day_04`, `ig_day_05`, `ig_day_06`, and `ig_day_07`, as shown below. 

Write a Python code that makes a **list** with the times spent on Instagram in sequence from day 1 to day 7. The list should be stored in a new variable `ig_list`.

```python
### Write your code below (in place of ...)
ig_day_01 = 120
ig_day_02 = 110
ig_day_03 = 15
ig_day_04 = 130
ig_day_05 = 135
ig_day_06 = 100
ig_day_07 = 120
# BEGIN SOLUTION
ig_list = [ig_day_01, ig_day_02, ig_day_03, ig_day_04, ig_day_05, ig_day_06, ig_day_07]
# END SOLUTION
```

```python
grader.check("q3a")
```

**b.** Immediately following the 7 consecutive days, the communication researcher continues to observe the time spent by the same individual on Instagram for three more consecutive days. On day 8, this individual spends 90 minutes on Instagram, on day 9 this individual spends 200 minutes on Instagram, and on day 10, they spend 145 minutes.

Make a new list which just has these three data points in sequence, i.e., time spent on day 8 comes first, then day 9, then day 10. 

Store your output in a new list called `ig_list_later`.

```python
### Write your code below (in place of ...)
ig_list_later = [90, 200, 145] # SOLUTION
```

```python
grader.check("q3b")
```

**c.** Make an updated list which combines the two lists you have just made, `ig_list` and `ig_list_later`, such that it contains the time spent by this individual on Instagram in sequence from day 1 to day 10.

Store your output in a new list called `ig_list_updated`.

**Tip:** You can add to lists to join them.

```python
### Write your code below (in place of ...)
ig_list_updated = ig_list + ig_list_later # SOLUTION
```

```python
grader.check("q3c")
```

**d.** Now, this communication researcher wants to know on which of the 10 days this individual spent exactly 2 hours on Instagram. 

Make a new list, `ig_twohours`, which has 10 boolean values in sequence for each day, True if individual spent exactly 2 hours on that day, and False if not.

For example, if the individual spent exactly 2 hours on days 3 and 4 only, the list would look like:
`[False, False, True, True, False, False, False, False, False, False]`

**Tip:** The measurements being collected by this communication researcher are in minutes, not hours. \
**Another tip:** If you convert `ig_list_updated` to a numpy array, you can use "equals to" operator with the entire array in one go! 

**Important:** Regardless of whether you use numpy array or lists, the final output should be a **list**. If you are using numpy arrays, make sure you convert the numpy array to a list (you can typecast using `list()`). 

```python
### Write your code below (in place of ...)
# BEGIN SOLUTION
import numpy
ig_twohours = list(numpy.array(ig_list_updated) == 120)
# END SOLUTION
```

```python
grader.check("q3d")
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Assignment #2 for COMM187: Data Science in Communication Research. Once finished answering questions, first **SAVE** then download this .ipynb file. Submit the file as instructed on Canvas and Gradescope. **ONLY** submit the .ipynb file, not the pdf.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True)
```

Tests for q1 (points: 1)
```
>>> bool_and == False
True
```
```
>>> bool_or == True
True
```
```
>>> bool_xor == True
True
```


Tests for q2 (points: 3)
```
>>> ismorethan500_indv_01 == False
True
```
```
>>> ismorethan500_indv_02 == False
True
```
```
>>> ismorethan500_indv_01 == False
True
```


Tests for q3a (points: 4)
```
>>> len(ig_list) == 7
True
```
```
>>> ig_list == [120, 110, 15, 130, 135, 100, 120]
True
```


Tests for q3b (points: 4)
```
>>> len(ig_list_later) == 3
True
```
```
>>> ig_list_later == [90, 200, 145]
True
```


Tests for q3c (points: 2)
```
>>> len(ig_list_updated) == 10
True
```
```
>>> ig_list_updated == [120, 110, 15, 130, 135, 100, 120, 90, 200, 145]
True
```


Tests for q3d (points: 6)
```
>>> len(ig_twohours) == 10
True
```
```
>>> ig_twohours == [True, False, False, False, False, False, True, False, False, False]
True
```
