```python
# Initialize Otter
import otter
grader = otter.Notebook("Assignment3_COMM187Spring2024.ipynb")
```

# Coding Assignment #3
## COMM187: Data Science in Communication Research, Spring 2024

**!!! Please make sure to run the first cell before running auto-grader !!!**

<span style="color:red">**BEFORE STARTING WITH THE ASSIGNMENT QUESTIONS, GO THROUGH THESE PRACTICE CODES AND NEW FUNCTIONS FIRST**</span>

### PRACTICE: Functions in `numpy` to create arrays

In Lab #2, go through the section on "Functions in `numpy` to create arrays" and practice the following functions in `numpy`:

1. `numpy.ones(n)`: produces an array of length `n` of all `1`s.
2. `numpy.zeros(n)`: produces an array of length `n` of all `0`s.

Now, try out these functions, for `n = 10`. In other words, make a numpy array with 10 `1`'s, and a numpy array with `0 `0`'s. 

```python
import numpy
numpy.ones(10)
```

Output:
```
array([1., 1., 1., 1., 1., 1., 1., 1., 1., 1.])
```

```python
numpy.zeros(10)
```

Output:
```
array([0., 0., 0., 0., 0., 0., 0., 0., 0., 0.])
```

### NEW FUNCTIONS: `size`, `min`, `max`, `sum`, and `mean`

Familiarize yourself with the following functions which can be operated only on `numpy` arrays.

**1.** `numpy.ndarray.size`

Just add `.size` at the end of any numpy array variable name, and this function will give you the number of values in that array.

NOTE: There is no need for the typical `()` brackets in this function.

```python
numpy_array1 = numpy.ones(100)
numpy_array1.size
```

Output:
```
100
```

**2.** `numpy.ndarray.min()`

Just add `.min()` at the end of any numpy array variable name, and this function will give you the minimum value in the numpy array.

```python
numpy_array2 = numpy.array([-1, 5, -10])
numpy_array2.min()
```

Output:
```
-10
```

**3.** `numpy.ndarray.max()`

Just add `.max()` at the end of any numpy array variable name, and this function will give you the minimum value in the numpy array.

```python
numpy_array3 = numpy.array([-1, 5, -10])
numpy_array3.max()
```

Output:
```
5
```

**4.** `numpy.ndarray.sum()`

Just add `.sum()` at the end of any numpy array variable name, and this function will calculate the total arithmetic sum of all the values in the numpy array. Basically, it will add up all the values.

```python
numpy_array4 = numpy.array([1,2,3])
numpy_array4.sum()
```

Output:
```
6
```

**5.** `numpy.ndarray.mean()`

Just add `.mean()` at the end of your numpy array variable name, and this function will calculate the mean or average of the numpy array.  

```python
numpy_array5 = numpy.array([1,2,3,4,5])
numpy_array5.mean()
```

Output:
```
3.0
```

After practicing these functions, solve the following questions.

---

### Health Communication Data Analysis

You are a data analyst working on a health communication research project about persuading people to be more physically active. 

Towards this, the project has collected daily number of minutes of light, moderate, and vigorous physical activity various individuals over a period of two years. Your task is to analyze this data using numpy arrays and perform several operations to gain insights into the data.

**Data:** For a single individual in the study sample, the data for their daily number of minutes of light, moderate, and vigorous physical activity are stored in the following variables:

`light` is a numpy array with the daily number of minutes spent doing light physical activity. \
`moderate` is a numpy array with the daily number of minutes spent doing moderate physical activity. \
`vigorous` is a numpy array with the daily number of minutes spent doing moderate physical activity. \ 

In each array, the first position contains minutes of activity for the first day, second position contains minutes of activity for the second day, so on and so forth for consecutive days over the period of two years.

Run the following cell before moving on the assignment questions.

```python
import numpy
light = numpy.loadtxt("./data/light.txt", delimiter = ",", dtype = int)
moderate = numpy.loadtxt("./data/moderate.txt", delimiter = ",", dtype = int)
vigorous = numpy.loadtxt("./data/vigorous.txt", delimiter = ",", dtype = int)
```

**Question 1: Basic Description of Data**


**a.** 
The data for this individual has been collected on a daily basis for two years, and stored in the numpy array variables mentioned above. What is the length of these numpy arrays?

Store the length of `light` in a new variable `light_length`. \
Store the length of `moderate` in a new variable `moderate_length`. \
Store the length of `vigorous` in a new variable `vigorous_length`.

*Note:* The data is collected every day for two consecutive years, assuming these are non-leap years.

```python
### Write your code below (in place of ...)
light_length = light.size # SOLUTION
moderate_length = moderate.size # SOLUTION
vigorous_length = vigorous.size # SOLUTION
```

```python
grader.check("q1a")
```

Output:
```
q1a results: All test cases passed!
```

**b.** For this individual, calculate the total number of minutes of physical activity of each type (light, moderate, and vigorous) in the last two years.

Store the total minutes of activity for light activity in a new variable `light_total`. \
Store the total minutes of activity for moderate activity in a new variable `moderate_total`. \
Store the total minutes of activity for vigorous activity in a new variable `vigorous_total`.

```python
### Write your code below (in place of ...)
light_total = light.sum() # SOLUTION
moderate_total = moderate.sum() # SOLUTION
vigorous_total = vigorous.sum() # SOLUTION
```

```python
grader.check("q1b")
```

Output:
```
q1b results: All test cases passed!
```

**c.** Now, calculate the **daily average number of minutes** of physical activity of each type (light, moderate, and vigorous) in the last two years for this individual.

Store the daily average minutes of activity for light activity in a new variable `light_avg`. \
Store the daily average minutes of activity for moderate activity in a new variable `moderate_avg`. \
Store the daily average  minutes of activity for vigorous activity in a new variable `vigorous_avg`.

```python
### Write your code below (in place of ...)
light_avg = light.mean() # SOLUTION
moderate_avg = moderate.mean() # SOLUTION
vigorous_avg = vigorous.mean() # SOLUTION
```

```python
grader.check("q1c")
```

Output:
```
q1c results: All test cases passed!
```

**Question 2: Data Correction** 

It turns out that there were some errors in measurement! The total number of minutes measured on the following days : \
Day 46, \
Day 105, \
Day 341, and \
Day 620 \
were 4 minutes less than they should have measured. 

This applies to light, moderate, and vigorous activities. Thus, for each of these days, we need to add 4 minutes to the existing data.

Remember, for each array `light`, `moderate`, and `vigorous`, the data in the first position is the data for Day 1, the data in the second position is data for Day 2, and so on and so forth.

Let us break this task down into smaller steps.

**a. Step 1:** Create an numpy array of the same size as that of `light`, `moderate`, and `vigorous`  with only `0`'s, and store this into a new variable named `corrections`.

```python
### Write your code below (in place of ...)
corrections = numpy.zeros(light.size) # SOLUTION
```

```python
grader.check("q2a")
```

Output:
```
q2a results: All test cases passed!
```

**b. Step 2:** In the numpy array `corrections`, change the value of the aforementioned positions to `4` instead of `0`.

*Hint example:* To change the value of the first element of a numpy array named `numpy_array` to a new value, let's say, `1`, you would need to run the following line of code: \
`numpy_array[0] = 1`

```python
### Write your code below (in place of ...)
# BEGIN SOLUTION
corrections[45] = 4
corrections[104] = 4
corrections[340] = 4
corrections[619] = 4
# END SOLUTION
```

```python
grader.check("q2b")
```

Output:
```
q2b results: All test cases passed!
```

**c. Step 3:** Great job so far! Now, in order to make the error correction, you just need to add `corrections` to each of the original arrays `light`, `moderate`, and `vigorous`, and save the output in a new array each.

Store the updated data for light activity in a new variable `light_updated`. \
Store the updated data for light activity in a new variable `moderate_updated`. \
Store the updated data for light activity in a new variable `vigorous_updated`. 

**Note:** Do NOT alter the data in the original arrays!

```python
### Write your code below (in place of ...)
light_updated = light + corrections # SOLUTION
moderate_updated = moderate + corrections # SOLUTION
vigorous_updated = vigorous + corrections # SOLUTION
```

```python
grader.check("q2c")
```

Output:
```
q2c results: All test cases passed!
```

**Question 3:** Now, with the updated numpy arrays, answer these analytical questions about the data.

**a.** In the last two years, for how many days did this individual perform light physical activity for more than an hour (60 minutes)?

Break this question down into two steps:

 - **Step 1:** Compare the values of `light_updated` with the threshold (one hour, or 60 minutes). This would look like: `light_updated > 60`. Look at the output. You should observe a numpy array of `bool` values.
 - **Step 2:** Now, store this output in a new variable `light_greaterthananhour`.
 - **Step 3:** Add the True values in `light_greaterthananhour` and store them in `light_greaterthananhour_total`. \
   Did you know, that if you simply use the `sum()` function (as described at the beginning of this assignment) on a boolean numpy array, it will simply add all the True values in that array? Use this to count the number of True values in `light_greaterthananhour`!

```python
### Write your code below (in place of ...)
light_greaterthananhour = light_updated > 60 # SOLUTION
light_greaterthananhour_total = light_greaterthananhour.sum() # SOLUTION
```

```python
grader.check("q3a")
```

Output:
```
q3a results: All test cases passed!
```

**b.** Similarly, in the last two years, for how many days did this individual perform moderate physical activity for more than half an hour (30 minutes)?

Break this question down into two steps just like before. The final answer should be stored in `moderate_morethanhalfhour_total`.

*Note:* Use `moderate_updated`, and not `moderate`!

```python
### Write your code below (in place of ...)
moderate_morethanhalfhour_total = (moderate_updated > 30).sum() # SOLUTION
```

```python
grader.check("q3b")
```

Output:
```
q3b results: All test cases passed!
```

**c.** Similarly, in the last two years, for how many days did this individual perform vigorous physical activity for more than 10 minutes?

Break this question down into two steps just like before. The final answer should be stored in `vigorous_morethan10mins_total`.

*Note:* Use `vigorous_updated`, and not `vigorous`!

```python
### Write your code below (in place of ...)
vigorous_morethan10mins_total = (vigorous_updated > 10).sum() # SOLUTION
```

```python
grader.check("q3c")
```

Output:
```
q3c results: All test cases passed!
```

**d.** Lastly, in the last two years, count the number of days on which this individual performed vigorous physical activity for longer amount of time than light activity. 

The final answer should be stored in a variable named `light_morethanvigorous_total`.

**Hint:** You can compare two same sized numpy arrays `array1` and `array2` by simply executing the following code: `array1 > array2` or `array1 < array2`.

*Note:* Use the updated arrays!

```python
### Write your code below (in place of ...)
light_morethanvigorous_total = (vigorous_updated > light_updated).sum() # SOLUTION
```

```python
grader.check("q3d")
```

Output:
```
q3d results: All test cases passed!
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Assignment #3 for COMM187: Data Science in Communication Research. Once finished answering questions, first **SAVE** then download this .ipynb file. Submit the file as instructed on Canvas and Gradescope. **ONLY** submit the .ipynb file, not the zip file.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True)
```

Output:
```
Running your submission against local test cases...


Your submission received the following results when run against available test cases:

    q1a results: All test cases passed!

    q1b results: All test cases passed!

    q1c results: All test cases passed!

    q2a results: All test cases passed!

    q2b results: All test cases passed!

    q2c results: All test cases passed!

    q3a results: All test cases passed!

    q3b results: All test cases passed!

    q3c results: All test cases passed!

    q3d results: All test cases passed!
<IPython.core.display.HTML object>
```

Tests for q1a (points: 3)
```
>>> light_length == 730
True
```
```
>>> moderate_length == 730
True
```
```
>>> vigorous_length == 730
True
```


Tests for q1b (points: 3)
```
>>> light_total == light.sum()
True
```
```
>>> moderate_total == moderate.sum()
True
```
```
>>> vigorous_total == vigorous.sum()
True
```


Tests for q1c (points: 3)
```
>>> light_avg == light.mean()
True
```
```
>>> light_avg - light.mean() < 0.01
True
```
```
>>> moderate_avg == moderate.mean()
True
```
```
>>> moderate_avg - moderate.mean() < 0.01
True
```
```
>>> vigorous_avg == vigorous.mean()
True
```
```
>>> vigorous_avg - vigorous.mean() < 0.01
True
```


Tests for q2a (points: 1)
```
>>> numpy.array_equal(corrections, numpy.zeros(light.size))
True
```


Tests for q2b (points: 2)
```
>>> corrections[45] == 4
True
```
```
>>> corrections[104] == 4
True
```
```
>>> corrections[340] == 4
True
```
```
>>> corrections[619] == 4
True
```
```
>>> temp_array = numpy.zeros(light.size)
>>> temp_array[45] = 4
>>> temp_array[104] = 4
>>> temp_array[340] = 4
>>> temp_array[619] = 4
>>> numpy.array_equal(corrections, temp_array)
True
```


Tests for q2c (points: 3)
```
>>> temp_array = numpy.zeros(light.size)
>>> temp_array[45] = 4
>>> temp_array[104] = 4
>>> temp_array[340] = 4
>>> temp_array[619] = 4
>>> numpy.array_equal(light_updated, light + temp_array)
True
```
```
>>> temp_array = numpy.zeros(light.size)
>>> temp_array[45] = 4
>>> temp_array[104] = 4
>>> temp_array[340] = 4
>>> temp_array[619] = 4
>>> numpy.array_equal(moderate_updated, moderate + temp_array)
True
```
```
>>> temp_array = numpy.zeros(light.size)
>>> temp_array[45] = 4
>>> temp_array[104] = 4
>>> temp_array[340] = 4
>>> temp_array[619] = 4
>>> numpy.array_equal(vigorous_updated, vigorous + temp_array)
True
```


Tests for q3a (points: 3)
```
>>> numpy.array_equal(light_greaterthananhour, light_updated > 60)
True
```
```
>>> light_greaterthananhour_total == (light_updated > 60).sum()
True
```


Tests for q3b (points: 3)
```
>>> moderate_morethanhalfhour_total == (moderate_updated > 30).sum()
True
```


Tests for q3c (points: 3)
```
>>> vigorous_morethan10mins_total == (vigorous_updated > 10).sum()
True
```


Tests for q3d (points: 3)
```
>>> light_morethanvigorous_total == (vigorous_updated > light_updated).sum()
True
```
